Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

NAME = Nicole Paraschiv
COLLABORATORS = ""

---

# CSCI 3155 Assignment 2 : Recursion and Tail Recursion

This assignment asks you to write scala programs. 

**Restrictions** apply to each problem in terms of forbidden Scala features and API functions. Please read them carefully and ask for clarifications from the course staff over Piazza or during office hours if unsure.

Note: `???` indicates that there is a missing function or code fragment that needs to be filled in. In scala, 
it is also a macro that throws a `NotImplemented` exception. Make sure that you remove the `???` and replace it with the answer. 

Use the test cases provided to test them. You are also encouraged to write your own test cases to help debug your work. However, please delete any extra cells you may have created lest they break our autograder.

**Very Important:** 
1. Please run the cell that defines the functions `passed` and `testWithMessage` below whenever you restart the notebook.
2. Use `@tailrec` annotation wherever a tail recursive function is expected. 

### Your Name Here

```scala
// TEST HELPER

// FIRST RUN THIS CELL EVERY TIME YOU START THE NOTEBOOK
def passed(points: Int) = {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}

def testWithMessage[T](v1: T, expected: T, testID: String) = { 
    println(s"Test $testID"); 
    println(s"\t Expected: $expected, your code returned: $v1")
    assert (v1 == expected, s"Test $testID FAILED.")
    println("\t Passed!")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
defined [32mfunction[39m [36mtestWithMessage[39m
```

## Problem 1A  (5 points)

Write a function `countEven` that inputs a list of numbers `lst` and counts how many elements of `lst` are even. 


### Restrictions 
 - Your function must be recursive (not necessary that it be tail recursive)
 - You __cannot__ use loops (for-loops, while loops etc..), and  __no__ mutables (var).
 - You __cannot__ use return statement in your function.
 - List Operations allowed: 
    - list.length - length of a list
    - list.head - extract the first element of a list
    - list.tail - extract the sublist from the second element  to last element of  list.
 - No other list API functions or API functions of any other datastructure allowed. Eg., do not look to convert your list into an array and use some Array API function.

<!-- grade_id: cell-0cb565560e29312d, points: 0, grade: false, solution: true -->
```scala
// YOUR CODE HERE
def countEven (lst: List[Int]): Int ={ 
    if(lst == List()) 0 
    else if (lst.head % 2 == 0) 1 + countEven(lst.tail) 
    else countEven(lst.tail)
}
```

Output:
```
defined [32mfunction[39m [36mcountEven[39m
```

<!-- grade_id: cell-c6190c2906bcedea, points: 5, grade: true, solution: false -->
```scala
testWithMessage( countEven(List(2, 1, 4, 1, 0, 1)), 3 , " countEven(List(0, 1, 0, 1, 0, 1)) # 1")
testWithMessage( countEven(List(0)), 1, " countEven(List(0)) # 2")
testWithMessage( countEven(List(2, 6, 4, 8)), 4, " countEven(List(2, 3, 4, -2)) # 3")
testWithMessage( countEven(List(5, 1, 2, 3, 4, 5, 0, 1, 5)), 3, " countEven(List(5, 1, 2, 3, 4, 5, 0, 1, 5)) # 5")
testWithMessage( countEven(List()), 0, "countEven(List()) # 6")
passed(5)
```

Output:
```
Test  countEven(List(0, 1, 0, 1, 0, 1)) # 1
	 Expected: 3, your code returned: 3
	 Passed!
Test  countEven(List(0)) # 2
	 Expected: 1, your code returned: 1
	 Passed!
Test  countEven(List(2, 3, 4, -2)) # 3
	 Expected: 4, your code returned: 4
	 Passed!
Test  countEven(List(5, 1, 2, 3, 4, 5, 0, 1, 5)) # 5
	 Expected: 3, your code returned: 3
	 Passed!
Test countEven(List()) # 6
	 Expected: 0, your code returned: 0
	 Passed!

*** Tests Passed (5 points) ***
```

## Problem 1B (5 points)

Write a tail recursive version of the `countEven` function  described in Problem 1A. Call your tail recursive function `countEvenTail`. The tests call `countEvenTail` with only one argument (the list). However, the function may take additional additional arguments (such as`acc`) as long as the additional arguments have a default value so that the tests wont fail. As discussed in the class, you may want to write a loop-based solution and use it to guide your tail recursive solution.

Place a `@tailrec` decorator in font of your function. 
Some scala versions (Scala < 2.12?) may require you to declare your function with the `final` keyword in front. 

~~~
final def countEvenTail(lst:List[Int],....): Int = { ... }
~~~


### Restrictions

Same restrictions as problem 1A.

Additionally your function must be tail recursive.

Calling `countEven` from inside your `countEvenTail` function is not allowed.

<!-- grade_id: cell-7326b99db5f0bc73, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec


@tailrec
// YOUR CODE HERE
def countEvenTail(lst: List[Int], runningSum: Int = 0): Int ={ 
    if(lst == List()) runningSum 
    else if (lst.head % 2 == 0) countEvenTail(lst.tail, runningSum+1) 
    else countEvenTail(lst.tail, runningSum)
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec


[39m
defined [32mfunction[39m [36mcountEvenTail[39m
```

<!-- grade_id: cell-96da0d4bea949b29, points: 5, grade: true, solution: false -->
```scala
testWithMessage( countEvenTail(List(2, 1, 4, 1, 0, 1)), 3 , " countEvenTail(List(0, 1, 0, 1, 0, 1)) # 1")
testWithMessage( countEvenTail(List(0)), 1, " countEvenTail(List(0)) # 2")
testWithMessage( countEvenTail(List(2, 6, 4, 8)), 4, " countEvenTail(List(2, 3, 4, -2)) # 3")
testWithMessage( countEvenTail(List(5, 1, 2, 3, 4, 5, 0, 1, 5)), 3, " countEvenTail(List(5, 1, 2, 3, 4, 5, 0, 1, 5)) # 5")
testWithMessage( countEvenTail(List()), 0, "countEvenTail(List()) # 6")
passed(5)
```

Output:
```
Test  countEvenTail(List(0, 1, 0, 1, 0, 1)) # 1
	 Expected: 3, your code returned: 3
	 Passed!
Test  countEvenTail(List(0)) # 2
	 Expected: 1, your code returned: 1
	 Passed!
Test  countEvenTail(List(2, 3, 4, -2)) # 3
	 Expected: 4, your code returned: 4
	 Passed!
Test  countEvenTail(List(5, 1, 2, 3, 4, 5, 0, 1, 5)) # 5
	 Expected: 3, your code returned: 3
	 Passed!
Test countEvenTail(List()) # 6
	 Expected: 0, your code returned: 0
	 Passed!

*** Tests Passed (5 points) ***
```

## Problem 1 C (10 points)

Write a tail recursive function `extractSubList` that inputs three arguments: a list `lst` of integers and indices `i` and `j`.

Given a list of integers `lst` and indices `i` and `j` into the list, extract all the elements in the original list `lst(i), ..., lst(j)` into a new list. Handle corner cases by returning an empty list when (a) `j < i` or (b) `j` is greater than or equal to the length of the list `lst`.

__Hint:__ Write a function `extractSubList` that handles all the _corner_ cases above and calls a "helper" function `extractSubListHelper` for all the _non corner_ cases. Make sure that `extractSubListHelper` is itself tail recursive. Make sure that the `@tailrec` decorator is applied to the helper function.

### Restrictions

  - Same restrictions as problem 1A.
  - Additionally your helper function must be tail recursive.
  - You __can__ use `list.head` and `list.tail` functions. Read about them here: https://www.tutorialspoint.com/scala/scala_lists.htm
  - You __can__ use `list.length`
  - You __can__ use list cons operator `elt::list` to append an element in front of a list, and also use `:+` operator to append to the back of the list. You __can__ also use `++` operator (or `:::` operator) to append two lists to each other.
  - You __can__ reverse a list by calling the List API function `list.reverse`.
  - You __cannot__ access list elements using their indices: Eg., `lst(i)` to directly access the ith element is forbidden.
  - No other list API functions allowed.

<!-- grade_id: cell-2ea676480d0a3998, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec


@tailrec
// YOUR CODE HERE
def extractSubListHelper(l: List[Int], k: Int, j: Int, i:Int, b: List[Int]): List[Int] = {
    if(k>j) b.reverse
    else if (k < i) extractSubListHelper(l.tail, k+1, j, i, b)
    else extractSubListHelper(l.tail, k+1,j, i, l.head :: b)
}
def extractSubList(l: List[Int], i: Int, j: Int): List[Int] ={
    if(j < i) List()
    else if (j >= l.length) List()
    else extractSubListHelper(l, 0,j, i, List())
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec


[39m
defined [32mfunction[39m [36mextractSubListHelper[39m
defined [32mfunction[39m [36mextractSubList[39m
```

<!-- grade_id: cell-ea11f3b0e46ecd37, points: 10, grade: true, solution: false -->
```scala
//BEGIN TESTS
testWithMessage(
    extractSubList(List(1, 2, 3, 4, 5), 2, 4),
    List(3, 4, 5),
    "test1"
    )

testWithMessage(
    extractSubList(List(1, 2, 3, 4, 5), 3, 5),
    List(),
    "test2"
    )

testWithMessage(
    extractSubList(List(1, 2, 3, 4, 5), 6, 5),
    List(),
    "test3"
    )

testWithMessage(
    extractSubList(List(1, 2, 3, 4, 5), 1, 1),
    List(2),
    "test4"
    )

testWithMessage(
    extractSubList(List(), 0, 0),
    List(),
    "test5"
    )

testWithMessage(
    extractSubList(List(1, 2, 3, 4, 5), 0, 4),
    List(1, 2, 3, 4, 5),
    "test6"
    )
passed(10)
//END TESTS
```

Output:
```
Test test1
	 Expected: List(3, 4, 5), your code returned: List(3, 4, 5)
	 Passed!
Test test2
	 Expected: List(), your code returned: List()
	 Passed!
Test test3
	 Expected: List(), your code returned: List()
	 Passed!
Test test4
	 Expected: List(2), your code returned: List(2)
	 Passed!
Test test5
	 Expected: List(), your code returned: List()
	 Passed!
Test test6
	 Expected: List(1, 2, 3, 4, 5), your code returned: List(1, 2, 3, 4, 5)
	 Passed!

*** Tests Passed (10 points) ***
```

## Problem 2 A (5 points)

Write a function `countVowels` that inputs a string `str` and counts the number of vowels `a, e, i, o, u, A, E, I, O, U` in the string. 

### Restrictions
 - Must be tail recursive.
 - You __cannot__ use loops (for-loops, while loops etc..), and  __no__ mutables (var).
 - You __cannot__ use return statement in your function.
 - String Operations allowed: 
    - `string.length` - length of a string
    - `string.head` - the first character in a string. Equivalent to `string(0)`
    - `string.tail` - extract the substring from the second character to end of string. Equivalent to `string.substring(1, string.length)`
    - No other string API functions are allowed.
 - Char Operations: 
    - `char.toLower` - Convert a char to lower case. Has no effect if char is already lower case or is not an alphabet.
 - No other list API functions or API functions of any other data structures allowed. Eg., do not look to convert your string into an array/list and use some Array/List API function.


<!-- grade_id: cell-2e7bb9392d8c5bd3, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec


@tailrec
// YOUR CODE HERE
def countVowels(str: String, countOfVowels: Int=0): Int = {
    if(str.length == 0) {
        countOfVowels
    }
    else{
        val c = str.head.toLower
        if (c == 'a' || c == 'e' || c == 'i' || c == 'o' || c == 'u'){
            countVowels(str.tail, countOfVowels+1)
        }
        else{
            countVowels(str.tail, countOfVowels)
        }
    }
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec


[39m
defined [32mfunction[39m [36mcountVowels[39m
```

<!-- grade_id: cell-4ba3adff731ac05c, points: 5, grade: true, solution: false -->
```scala
//BEGIN TEST
testWithMessage( countVowels("atchou") , 3, "\"atchou\" must have three vowels in it" )
testWithMessage( countVowels(""), 0, "Empty string has no vowels")
testWithMessage( countVowels("$%!%@!1125122^^&&"), 0, "Non alphabet chars handled")
testWithMessage( countVowels("AAEIOU"), 6, "Capitalized vowels handled")
testWithMessage( countVowels("Grzzly"), 0, "\"Grzzly\" has no vowels")
testWithMessage( countVowels("Meow, Meow, says the alley cat next door!!"), 12, "Handling commas and spaces?")
passed(5)
//END TEST
```

Output:
```
Test "atchou" must have three vowels in it
	 Expected: 3, your code returned: 3
	 Passed!
Test Empty string has no vowels
	 Expected: 0, your code returned: 0
	 Passed!
Test Non alphabet chars handled
	 Expected: 0, your code returned: 0
	 Passed!
Test Capitalized vowels handled
	 Expected: 6, your code returned: 6
	 Passed!
Test "Grzzly" has no vowels
	 Expected: 0, your code returned: 0
	 Passed!
Test Handling commas and spaces?
	 Expected: 12, your code returned: 12
	 Passed!

*** Tests Passed (5 points) ***
```

## Problem 2B (10 points)

Write a function `extractPrefix` that takes in a string `str`, and a length `j`. It extracts the first `j` characters of the string if `j < str.length` or else returns the entire string if `j >= str.length`. If `j < 0`, then `IllegalArgumentException` must be thrown using `require`.

### Examples

  - `extractPrefix("Hello", 3)` should return `"Hel"`
  - `extractPrefix("Yallooo", 1)` should return `"Y"`
  - `extractPrefix("AnyongHaseyo", 0)` should return `""` (empty string)
  - `extractPrefix("hello", 10)` should return `"hello"` (entire string)
  - `extractPrefix("hello", -1)` should throw an `IllegalArgumentException`.
 
### Restrictions
- Must be tail recursive.
- You __cannot__ use loops (for-loops, while loops etc..), and  __no__ mutables (var).
- You __cannot__ use return statement in your function.
- String Operations allowed: 
    - `string.length` - length of a string
    - `string.head` - the first character in a string. Equivalent to `string(0)`
    - `string.tail` - extract the substring from the second character to end of string. Equivalent to `string.substring(1, string.length)`
    - Appending two strings `s1 + s2` is permitted. It is also permitted to append `s + c` where `s` is a string and `c` is a Char. This will yield a new string with the character concatenated  to the end of the string.
- No other list API functions or API functions of any other data structures allowed. Eg., do not look to convert your string into an array/list and use some Array/List API function.




<!-- grade_id: cell-d6dc9fd0bc67aa68, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec


@tailrec
// YOUR CODE HERE
def extractPrefix(str: String, j: Int, acc: String =""): String = {
    require(j >= 0, "IllegalArgumentException")
    if(j == 0 || 0 == str.length()) {
        acc
    }
    else{
        extractPrefix(str.tail,j-1, acc+str.head)
    }
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec


[39m
defined [32mfunction[39m [36mextractPrefix[39m
```

<!-- grade_id: cell-c350ba564b712141, points: 10, grade: true, solution: false -->
```scala
testWithMessage(extractPrefix("Hello", 3), "Hel", "Test # 1")
testWithMessage(extractPrefix("Hel", 6), "Hel", "Test # 2")
testWithMessage(extractPrefix("", 10), "", "Test # 3")
testWithMessage(extractPrefix("Papa", 0), "", "Test # 4")

testWithMessage(extractPrefix("Papa", 5), "Papa", "Test # 5")
testWithMessage(extractPrefix("Papa", 4), "Papa", "Test # 6")
try {
    testWithMessage(extractPrefix("Papa", -1), "IllegalArgumentException", "Test #7")
    
} catch {
    case e: IllegalArgumentException => 
      testWithMessage("IllegalArgumentException", "IllegalArgumentException", "Test # 7")
    case _: Throwable =>
      testWithMessage("Some other exception", "IllegalArgumentException", "Test # 7")
}

passed(10)
```

Output:
```
Test Test # 1
	 Expected: Hel, your code returned: Hel
	 Passed!
Test Test # 2
	 Expected: Hel, your code returned: Hel
	 Passed!
Test Test # 3
	 Expected: , your code returned: 
	 Passed!
Test Test # 4
	 Expected: , your code returned: 
	 Passed!
Test Test # 5
	 Expected: Papa, your code returned: Papa
	 Passed!
Test Test # 6
	 Expected: Papa, your code returned: Papa
	 Passed!
Test Test # 7
	 Expected: IllegalArgumentException, your code returned: IllegalArgumentException
	 Passed!

*** Tests Passed (10 points) ***
```

## Problem 3 (15 points)

The Newton-Raphson method for finding the square root of a positive number $s > 0$ works as follows:

  - Set $x_0 = s$ as the starting value.
  - At each step $n$, set $x_n = 0.5 ( x_{n-1} + \frac{s}{x_{n-1}}) $.
  - Terminate whenever $| x_n^2 - s| \leq \epsilon$, where $\epsilon $ is a small error tolerance (typically $10^{-10}$).

For your convenience we provide you with a while loop based version of this method to find square root. Your goal is to implement it as a tail recursive function.


```scala
def sqrtLoop(s: Double, epsilon: Double = 1E-10): Double = {
    require (s >= 0.0) // This is just a fancy assert statement that expresses a precondition
    var x:Double = s
    while ( math.abs(x*x - s) > epsilon) {
        x = 0.5 * (x + s/x)
    }
    x
}
```

Output:
```
defined [32mfunction[39m [36msqrtLoop[39m
```

```scala
/* -- Let's test our loopy version --*/
for (i <- 0 to 100){
    val x = i/10.0
    println(s"sqrt($x) = ${sqrtLoop(x)}" )
}
```

Output:
```
sqrt(0.0) = 0.0
sqrt(0.1) = 0.31622776601683794
sqrt(0.2) = 0.4472135954999956
sqrt(0.3) = 0.5477225575051661
sqrt(0.4) = 0.6324555320888272
sqrt(0.5) = 0.7071067811873449
sqrt(0.6) = 0.7745966692414905
sqrt(0.7) = 0.8366600265340756
sqrt(0.8) = 0.894427190999916
sqrt(0.9) = 0.9486832980509526
sqrt(1.0) = 1.0
sqrt(1.1) = 1.0488088481703692
sqrt(1.2) = 1.0954451150509241
sqrt(1.3) = 1.140175425099138
sqrt(1.4) = 1.1832159566199232
sqrt(1.5) = 1.2247448713915894
sqrt(1.6) = 1.2649110640673549
sqrt(1.7) = 1.3038404810405515
sqrt(1.8) = 1.3416407864999869
sqrt(1.9) = 1.3784048752094868
sqrt(2.0) = 1.4142135623746899
sqrt(2.1) = 1.4491376746236861
sqrt(2.2) = 1.4832396974316797
sqrt(2.3) = 1.516575088840445
sqrt(2.4) = 1.5491933384829668
sqrt(2.5) = 1.5811388300841895
sqrt(2.6) = 1.6124515496597098
sqrt(2.7) = 1.6431676725154984
sqrt(2.8) = 1.6733200530681511
sqrt(2.9) = 1.70293863659264
sqrt(3.0) = 1.7320508075688772
sqrt(3.1) = 1.760681686165901
sqrt(3.2) = 1.788854381999832
sqrt(3.3) = 1
```

Implement a tail recursive function `sqrtRec` which takes as an argument a double precision number `s` that is assumed to be non-negative, and optional argument `epsilon` which has a default value of `1E-10`.  You may implement a tail recursive _helper_ function and call that from your `sqrtRec` function.

### Restrictions

 - Use of any math function other than `math.abs` is disallowed. You cannot, for instance, use `math.sqrt` or `math.log`. However, you can use `math.abs` function.
 - You __cannot__ use loops (for-loops, while loops etc..), and  __no__ mutables (var).
 - You __cannot__ use return statement in your function.
 - Your functions must be tail recursive.
 - Use of any other data-structure/number API function is also disallowed.
 - If it is more convenient,  make a tail recursive function `sqrtRecHelper` and call that function from the desired `sqrtRec` function.

<!-- grade_id: cell-9793083fbebd61ef, points: 0, grade: false, solution: true -->
```scala
import scala.annotation.tailrec
@tailrec
// YOUR CODE HERE
def sqrtRecHelper(x: Double, s: Double, epsilon: Double = 1E-10): Double = {
    if(math.abs(x*x -s) <= epsilon) {
        x
    }
    else{
        sqrtRecHelper((0.5*(x+s/x)),s, epsilon)
    }
}
def sqrtRec(s: Double, epsilon: Double = 1E-10): Double = {
    require(s >=0.0)
    if(s == 0) 0
    else{
        sqrtRecHelper(s,s, epsilon)
    }
}
```

Output:
```
[32mimport [39m[36mscala.annotation.tailrec
[39m
defined [32mfunction[39m [36msqrtRecHelper[39m
defined [32mfunction[39m [36msqrtRec[39m
```

<!-- grade_id: cell-d296cb8a424fbeb2, points: 15, grade: true, solution: false -->
```scala
def sqrtLoop(s: Double, epsilon: Double = 1E-10): Double = {
    require (s >= 0.0) // This is just a fancy assert statement that expresses a precondition
    var x:Double = s
    while ( math.abs(x*x - s) > epsilon) {
        x = 0.5 * (x + s/x)
    }
    x
}

for (i <- 0 to 100){
    val x:Double = i.toDouble/10.0 + 0.05
    val y = sqrtRec(x, 1E-10)
    println(s"Computing sqrt($x)")
    if ( math.abs(y * y - x) > 1E-10){
        println(s"Your Implementation failed for s = $x, it returned $y, the expected answer is ${sqrtLoop(x)}" )
        assert (false)
    }
}
passed(15)
```

Output:
```
Computing sqrt(0.05)
Computing sqrt(0.15000000000000002)
Computing sqrt(0.25)
Computing sqrt(0.35)
Computing sqrt(0.45)
Computing sqrt(0.55)
Computing sqrt(0.65)
Computing sqrt(0.75)
Computing sqrt(0.8500000000000001)
Computing sqrt(0.9500000000000001)
Computing sqrt(1.05)
Computing sqrt(1.1500000000000001)
Computing sqrt(1.25)
Computing sqrt(1.35)
Computing sqrt(1.45)
Computing sqrt(1.55)
Computing sqrt(1.6500000000000001)
Computing sqrt(1.75)
Computing sqrt(1.85)
Computing sqrt(1.95)
Computing sqrt(2.05)
Computing sqrt(2.15)
Computing sqrt(2.25)
Computing sqrt(2.3499999999999996)
Computing sqrt(2.4499999999999997)
Computing sqrt(2.55)
Computing sqrt(2.65)
Computing sqrt(2.75)
Computing sqrt(2.8499999999999996)
Computing sqrt(2.9499999999999997)
Computing sqrt(3.05)
Computing sqrt(3.15)
Computing sqrt(3.25)
Computing sqrt(3.3499999999999996)
Computing sqrt(3.4499999999999997)
Computing sqrt(3.55)
Computing sqrt(3.65)
Computing sqrt(3.75)
Computing sqrt(3.8499999999999996)
Computing sq
```

## That's All Folks