Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

NAME = "Nicole Paraschiv"
COLLABORATORS = ""

---

# CSCI 3155 Spring 2026: Assignment 9

Goals of this assignment include: 
- Developing a type system for records and references

```scala
// TEST HELPER
def passed(points: Int) = {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}
```

Output:
```
defined [32mfunction[39m [36mpassed[39m
```

## Problem 1 (50 points)

In this problem, we will extend Lettuce language and its type system to records and mutable references.  Here is a stripped down version of Lettuce we are going to use in this problem.

$$\begin{array}{rcl}
\mathbf{Expr} & \rightarrow & Const(\mathbf{Double}) \\
& | & Ident(\mathbf{Identifier})\\
& | & Plus(\mathbf{Expr}, \mathbf{Expr})\\
& | & Div(\mathbf{Expr}, \mathbf{Expr}) \\
& | & Geq(\mathbf{Expr}, \mathbf{Expr}) \\
& | & And(\mathbf{Expr}, \mathbf{Expr}) \\
& | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr})\\
& | & Let(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}, \mathbf{Expr})\\
& | & FunDef(\mathbf{Identifier}, \color{red}{\mathbf{Type}}, \mathbf{Expr}) \\
& | & FunCall(\mathbf{Expr}, \mathbf{Expr}) \\
& | & \color{blue}{UnitConst} & \leftarrow \text{written as} \ () \\
& | & \color{blue}{NewRef(\mathbf{Expr})} \\
& | & \color{blue}{DeRef(\mathbf{Expr})} \\
& | & \color{blue}{AssignRef(\mathbf{Expr}, \mathbf{Expr})} & \leftarrow \text{written as} \ e_1 := e_2 \\
& | & \color{blue}{Record(List[(Identifier, \mathbf{Expr})])} & \leftarrow \text{written as} \ \{\overline{x=e}\}, \ \text{e.g.,} \ \{x_1=e_1, x_2=e_2, \ldots, x_i=e_i, \ldots \} \\
\end{array}$$

We extend the type system with a record type constructor. 
$$\begin{array}{rcl}
\mathbf{Type} & \rightarrow & NumType & \leftarrow \text{written as } \ num \\
& | & BoolType & \leftarrow \text{written as } \ bool \\
& | & UnitType & \leftarrow \text{written as } \ unit \\
& | & FunType(\mathbf{Type} , \mathbf{Type} ) & \leftarrow \text{written as } \ t_1 \Rightarrow t_2 \\
& | & \color{red}{RefType(\mathbf{Type})} & \leftarrow \text{written as } \ ref(t_1) \\
& | & \color{red}{RecordType(List[(Identifier, \mathbf{Type})])} & \leftarrow \text{written as} \ \{\overline{x:t}\}, \ \text{e.g.,} \ \{x_1:t_1, x_2:t_2, \ldots, x_i:t_i, \ldots \} \\  
\end{array}$$




### A (10 points; Manually Graded): Let us write some inference rules for working with records and ref types.

$$\newcommand\typeOf{\mathbf{typeOf}}$$
$$\newcommand\semRule[3]{\begin{array}{c} #1 \\ \hline #2 \\ \end{array}\;(\text{#3}) }$$
Recall from notes on "Types and Type Checking" that $\typeOf(\texttt{e}, \alpha)$ is the type of an 
expression $\texttt{e}$ under type environment $\alpha$. The type environment maps identifiers in the current scope to  their annotated types.

Let us write the type rule for the record expression. The rule is straightforward; it simply composes the types of all sub-expressions that make up the record.

$$\semRule{ \forall i.~\typeOf(e_i, \alpha) = t_i,\;\; t_i \not= \mathbf{typeerror} }
          {\typeOf(\{\overline{x=e}\}, \alpha) = \{\overline{x:t}\} }{record-ok}$$



Next, let us write a rule for NewRef.

$$\semRule{ \typeOf(\texttt{e}, \alpha) = t,\; t \not= \mathbf{typeerror} }{\typeOf(\texttt{NewRef(e)}, \alpha) = ref(t) }{newref-ok}$$

It says that if $\texttt{e}$ receives type $t$ under type environment $\alpha$ and it is not a type error, then $\texttt{NewRef(e)}$ must receive the type $ref(t)$ under $\alpha$.



(i) Complete the missing terms for the rule for `UnitConst` OK rule. The unit constant, i.e., `()`, always has the type `unit`.
$$\semRule{}{\typeOf((), \alpha) = \color{red}{???_1}}{unitconst-ok}$$

(ii) Complete the missing terms for the rule for `DeRef` OK rule. For example, dereferencing a reference to a number (`num`) should give you a number (`num`).
$$\semRule{\typeOf(\texttt{e}, \alpha) = ref(t)}{\typeOf(\texttt{DeRef(e)}, \alpha) = \color{red}{???_2}}{deref-ok}$$

(iii) Complete the missing terms for the rule for `DeRef` error.
$$\semRule{\typeOf(\texttt{e}, \alpha) \not= \color{red}{???_3} }{\typeOf(\texttt{DeRef(e)}, \alpha) = \mathbf{typeerror}}{deref-nok}$$

(iv) Complete the missing terms for `AssignRef` OK rule.
$$\semRule{\typeOf(\texttt{e1}, \alpha) = \color{red}{???_4}, \typeOf(\texttt{e2}, \alpha) = t}{\typeOf(\texttt{AssignRef(e1, e2)}, \alpha) =   unit}{assignref-ok}$$


In the cell below, write down the solution for each of the missing terms $\color{red}{???_1}$ - $\color{red}{???_4}$. Write each one on its own line.

<!-- grade_id: cell-b15cecf7fa15e6a0, points: 8, grade: true, solution: true -->
1. unit
2. t
3. ref(t)
4. ref(t)

### C (15 points): Write the type checker with references by filling in the missing code.

<!-- grade_id: cell-c8a290b7a2779282, points: 0, grade: false, solution: true -->
```scala
sealed trait Type
case object NumType extends Type
case object BoolType extends Type
case class FunType(t1: Type, t2: Type) extends Type
// TODO: Write new case classes for UnitType, RefType, and RecordType
// YOUR CODE HERE
case object UnitType extends Type
case class RefType(t: Type) extends Type
case class RecordType(ies: List[(String,Type)]) extends Type

sealed trait Program
sealed trait Expr

case class Const(f: Double) extends Expr
case class Ident(s: String) extends Expr
case class Plus(e1: Expr, e2: Expr) extends Expr
case class Minus(e1: Expr, e2: Expr) extends Expr
case class Geq(e1: Expr, e2: Expr) extends Expr
case class IfThenElse(e1: Expr, e2: Expr, e3: Expr) extends Expr
case class Let(x: String, xType: Type, e1: Expr, e2: Expr) extends Expr
case class FunDef(id: String, idType: Type, e: Expr) extends Expr
case class FunCall(calledFun: Expr, argExpr: Expr) extends Expr
case class NewRef(e: Expr) extends Expr
case class DeRef(e: Expr) extends Expr
case class AssignRef(e1: Expr, e2: Expr) extends Expr
case object UnitConst extends Expr
case class Record(ies: List[(String,Expr)]) extends Expr
case class TopLevel(e: Expr) extends Program

def typeEquals(t1: Type, t2: Type): Boolean = t1 == t2
case class TypeErrorException(s: String) extends Exception

```

Output:
```
defined [32mtrait[39m [36mType[39m
defined [32mobject[39m [36mNumType[39m
defined [32mobject[39m [36mBoolType[39m
defined [32mclass[39m [36mFunType[39m
defined [32mobject[39m [36mUnitType[39m
defined [32mclass[39m [36mRefType[39m
defined [32mclass[39m [36mRecordType[39m
defined [32mtrait[39m [36mProgram[39m
defined [32mtrait[39m [36mExpr[39m
defined [32mclass[39m [36mConst[39m
defined [32mclass[39m [36mIdent[39m
defined [32mclass[39m [36mPlus[39m
defined [32mclass[39m [36mMinus[39m
defined [32mclass[39m [36mGeq[39m
defined [32mclass[39m [36mIfThenElse[39m
defined [32mclass[39m [36mLet[39m
defined [32mclass[39m [36mFunDef[39m
defined [32mclass[39m [36mFunCall[39m
defined [32mclass[39m [36mNewRef[39m
defined [32mclass[39m [36mDeRef[39m
defined [32mclass[39m [36mAssignRef[39m
defined [32mobject[39m [36mUnitConst[39m
defined [32mclass[39m [36mRecord[39m
defined [32mclass[39m [36mTopLevel[39m
def
```

<!-- grade_id: cell-e3d9acd1877b00b4, points: 0, grade: false, solution: true -->
```scala
def typeOf(e: Expr, alpha: Map[String, Type]): Type = {
    def checkType(opName: String, e1: Expr, t1: Type, e2: Expr, t2: Type, resType: Type): Type = {
        val t1hat = typeOf(e1, alpha)
        if (! typeEquals(t1hat, t1)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t1, obtained $t1hat")
        }
        
        val t2hat = typeOf(e2, alpha)
        if (! typeEquals(t2hat, t2)){
            throw new TypeErrorException(s"Type mismatch in arithmetic/comparison/bool op $opName, Expected type $t2, obtained $t2hat")
        }
        
        resType
    }
    
    e match {
        case Const(f) => NumType
        case Ident(s) => {if (alpha contains s)
                             alpha(s)
                          else 
                             throw TypeErrorException(s"Unknown identifier $s")}
        case Plus(e1, e2) =>  checkType("Plus", e1,  NumType, e2, NumType, NumType)
        case Minus(e1, e2) => checkType("Minus",e1,  NumType, e2, NumType, NumType)
        case Geq(e1, e2) => checkType("Geq", e1,  NumType, e2, NumType, BoolType)
        case IfThenElse(e, e1, e2) => {
            val t = typeOf(e, alpha)
            if (t == BoolType){
                val t1 = typeOf(e1, alpha)
                val t2 = typeOf(e2, alpha)
                if (typeEquals(t1, t2))
                    t1
                else 
                    throw TypeErrorException(s"If then else returns unequal types $t1 and $t2")
            } else {
                throw TypeErrorException(s"If then else condition expression not boolean $t")
            }
        }

        case Let(x, t, e1, e2) => {
            val t1 = typeOf(e1, alpha)
            if (typeEquals(t1, t)){
                val newAlpha = alpha + (x -> t)
                typeOf(e2, newAlpha)
            } else {
                throw TypeErrorException(s"Let binding has type $t whereas it is bound to expression of type $t1")
            }
        }

        case FunDef(x, t1, e) => {
            val newAlpha = alpha + (x -> t1)
            val t2 = typeOf(e, newAlpha)
            FunType(t1, t2)
        }

        case FunCall(e1, e2) => {
            val ftype = typeOf(e1, alpha)
            ftype match {
                case FunType(t1, t2) => {
                    val argType = typeOf(e2, alpha)
                    if (typeEquals(argType, t1)){
                        t2
                    } else {
                        throw TypeErrorException(s"Call to function with incompatible argument type. Expected $t1, obtained $argType")
                    }
                }
                case _ => { throw TypeErrorException(s"Call to function but with a non function type $ftype")}

            }
        }

        case NewRef(e) => {
            // YOUR CODE HERE
            val argType = typeOf(e, alpha)
            RefType(argType)
        }
        
        case AssignRef(e1, e2) => {
            // YOUR CODE HERE
            val argType = typeOf(e1, alpha)
            val argT = typeOf(e2, alpha)
            argType match {
                case RefType(t) => {
                    if (typeEquals(t, argT)) {UnitType}
                    else {throw TypeErrorException(s"Error")}
                }
                case _ => throw TypeErrorException(s"Error")
                }
        }
        
        case DeRef(e) => {
            // YOUR CODE HERE
            val argType = typeOf(e, alpha)
            argType match {
                case RefType(t) => t
                case _ =>  throw TypeErrorException(s"Error")
            }
        }
        
        case Record(ies) => {
            // YOUR CODE HERE
            val typed = ies.map {
                case(id, expr) => (id, typeOf(expr,alpha))
            }
            RecordType(typed)
        }
    }
}

def typeOfProgram(p: Program) = p match {
    case TopLevel(e) => {
            val t = typeOf(e, Map())
            println(s"Program type computed successfully as $t")
            t
    }
}
```

Output:
```
-- [E029] Pattern Match Exhaustivity Warning: cmd6.sc:16:4 ---------------------
16 |    e match {
   |    ^
   |    match may not be exhaustive.
   |
   |    It would fail on pattern case: UnitConst
   |
   | longer explanation available when compiling with `-explain`
defined [32mfunction[39m [36mtypeOf[39m
defined [32mfunction[39m [36mtypeOfProgram[39m
```

<!-- grade_id: cell-f9cfa6cb83f4b9d0, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST

/* 
let x : ref(num) = NewRef(10 ) in 
   let dummy: unit = AssignRef(x, 30) in 
       DeRef(x)
       */

val p1 = Let("x", RefType(NumType), NewRef(Const(10)), Let("dummy", UnitType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t1 = typeOfProgram(TopLevel(p1))
assert(t1 == NumType, "Test 1 failed: answer should be NumType")
passed(6)
//END TEST
```

Output:
```
Program type computed successfully as NumType

*** Tests Passed (6 points) ***
[36mp1[39m: [32mLet[39m = [33mLet[39m(
  x = [32m"x"[39m,
  xType = [33mRefType[39m(t = NumType),
  e1 = [33mNewRef[39m(e = [33mConst[39m(f = [32m10.0[39m)),
  e2 = [33mLet[39m(
    x = [32m"dummy"[39m,
    xType = UnitType,
    e1 = [33mAssignRef[39m(e1 = [33mIdent[39m(s = [32m"x"[39m), e2 = [33mConst[39m(f = [32m30.0[39m)),
    e2 = [33mDeRef[39m(e = [33mIdent[39m(s = [32m"x"[39m))
  )
)
[36mt1[39m: [32mType[39m = NumType
```

<!-- grade_id: cell-e6d228f96bb27e62, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num) = NewRef(function(z: num) z + 10) in 
   let dummy: unit = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p2 = Let("x", RefType(NumType), NewRef(fdef), Let("dummy", UnitType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t2 = try {
   typeOfProgram(TopLevel(p2))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(6)
//END TEST
```

<!-- grade_id: cell-a38f35932eb5f24f, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: unit = AssignRef(NewRef(35), 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p4 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", UnitType, AssignRef(NewRef(Const(35)), Const(30) ), DeRef(Ident("x"))) )
val t4 =  typeOfProgram(TopLevel(p4))
assert(t4 == FunType(NumType, NumType), "Test failed")
passed(6)
//END TEST
```

Output:
```
Program type computed successfully as FunType(NumType,NumType)

*** Tests Passed (6 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m(
  id = [32m"z"[39m,
  idType = NumType,
  e = [33mPlus[39m(e1 = [33mIdent[39m(s = [32m"z"[39m), e2 = [33mConst[39m(f = [32m10.0[39m))
)
[36mp4[39m: [32mLet[39m = [33mLet[39m(
  x = [32m"x"[39m,
  xType = [33mRefType[39m(t = [33mFunType[39m(t1 = NumType, t2 = NumType)),
  e1 = [33mNewRef[39m(
    e = [33mFunDef[39m(
      id = [32m"z"[39m,
      idType = NumType,
      e = [33mPlus[39m(e1 = [33mIdent[39m(s = [32m"z"[39m), e2 = [33mConst[39m(f = [32m10.0[39m))
    )
  ),
  e2 = [33mLet[39m(
    x = [32m"dummy"[39m,
    xType = UnitType,
    e1 = [33mAssignRef[39m(e1 = [33mNewRef[39m(e = [33mConst[39m(f = [32m35.0[39m)), e2 = [33mConst[39m(f = [32m30.0[39m)),
    e2 = [33mDeRef[39m(e = [33mIdent[39m(s = [32m"x"[39m))
  )
)
[36mt4[39m: [32mType[39m = [33mFunType[39m(t1 = NumTy
```

<!-- grade_id: cell-9000c2d748b5c904, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
/* 
let x : ref(num => num) = NewRef(function(z: num) z + 10) in 
   let dummy: num = AssignRef(x, 30) in 
       DeRef(x)
       */
val fdef = FunDef("z", NumType, Plus(Ident("z"), Const(10)))
val p3 = Let("x", RefType(FunType(NumType, NumType)), NewRef(fdef), Let("dummy", NumType, AssignRef(Ident("x"), Const(30) ), DeRef(Ident("x"))) )
val t3 = try {
   typeOfProgram(TopLevel(p3))
   assert(false, "The program should not receive a type")
} catch {
    case TypeErrorException(msg) => s"OK -- caught a type error exception: $msg"
    case e => print(e); assert(false, "Please throw TypeErrorException(message) when a type failure occurs")
}
passed(6)
//END TEST
```

Output:
```
*** Tests Passed (6 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m(
  id = [32m"z"[39m,
  idType = NumType,
  e = [33mPlus[39m(e1 = [33mIdent[39m(s = [32m"z"[39m), e2 = [33mConst[39m(f = [32m10.0[39m))
)
[36mp3[39m: [32mLet[39m = [33mLet[39m(
  x = [32m"x"[39m,
  xType = [33mRefType[39m(t = [33mFunType[39m(t1 = NumType, t2 = NumType)),
  e1 = [33mNewRef[39m(
    e = [33mFunDef[39m(
      id = [32m"z"[39m,
      idType = NumType,
      e = [33mPlus[39m(e1 = [33mIdent[39m(s = [32m"z"[39m), e2 = [33mConst[39m(f = [32m10.0[39m))
    )
  ),
  e2 = [33mLet[39m(
    x = [32m"dummy"[39m,
    xType = NumType,
    e1 = [33mAssignRef[39m(e1 = [33mIdent[39m(s = [32m"x"[39m), e2 = [33mConst[39m(f = [32m30.0[39m)),
    e2 = [33mDeRef[39m(e = [33mIdent[39m(s = [32m"x"[39m))
  )
)
[36mt3[39m: [32mString[39m = [32m"OK -- caught a type error exception: Error"[39m
```

<!-- grade_id: cell-2f222c52168dcd0d, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST

val p1 = ("foo",Const(10.0))
val t1 = ("foo",NumType)

val p2 = ("bar",Geq(Const(9.0), Const(10.0)))
val t2 = ("bar",BoolType)

val rdec1 = Record(List(p1,p2))
val rtyp1 = RecordType(List(t1,t2))

val prog1 = TopLevel(rdec1)
assert(typeOfProgram(prog1) == rtyp1, "Test failed")
passed(2)

val p3 = ("baz",NewRef(Const(11.0)))
val t3 = ("baz",RefType(NumType))
val p4 = ("qux",rdec1)
val t4 = ("qux", rtyp1)

val rdec2 = Record(List(p3,p4))
val rtyp2 = RecordType(List(t3,t4))
// prog2 = {baz=11.0, qux={foo=10.0, bar=9.0>10.0}}
val prog2 = TopLevel(rdec2)
assert(typeOfProgram(prog2) == rtyp2, "Test failed")
passed(6)


//END TEST
```

Output:
```
Program type computed successfully as RecordType(List((foo,NumType), (bar,BoolType)))

*** Tests Passed (2 points) ***
Program type computed successfully as RecordType(List((baz,RefType(NumType)), (qux,RecordType(List((foo,NumType), (bar,BoolType))))))

*** Tests Passed (6 points) ***
[36mp1[39m: ([32mString[39m, [32mConst[39m) = ([32m"foo"[39m, [33mConst[39m(f = [32m10.0[39m))
[36mt1[39m: ([32mString[39m, scala.Tuple2[java.lang.String, cmd10.this.cmd5.NumType]) = (
  [32m"foo"[39m,
  NumType
)
[36mp2[39m: ([32mString[39m, [32mGeq[39m) = ([32m"bar"[39m, [33mGeq[39m(e1 = [33mConst[39m(f = [32m9.0[39m), e2 = [33mConst[39m(f = [32m10.0[39m)))
[36mt2[39m: ([32mString[39m, scala.Tuple2[java.lang.String, cmd10.this.cmd5.BoolType]) = (
  [32m"bar"[39m,
  BoolType
)
[36mrdec1[39m: [32mRecord[39m = [33mRecord[39m(
  ies = [33mList[39m(
    ([32m"foo"[39m, [33mConst[39m(f = [32m10.0[39m)),
    ([32m"bar"[39m, [33mGeq[39m(e1 = [33mConst
```

<!-- grade_id: cell-073d515f61d32752, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
/*
let x: ref(Num) = NewRef(0) in 
 let f : num => unit = function (z:num) { x := z} in 
   let obj: {set_counter:ref(num => unit)} = {set_counter = NewRef(f)} in
     obj
*/
val fdef = FunDef("z", NumType, AssignRef(Ident("x"), Ident("z")))
val gdef = FunDef("z", UnitType, DeRef(Ident("x")))
val recrd = Record(List(("set_counter",NewRef(Ident("f")))))
val rectyp = RecordType(List(("set_counter",RefType(FunType(NumType,UnitType)))))
val p3 = Let("x", RefType(NumType), NewRef(Const(0.0)), 
             Let("f", FunType(NumType, UnitType), fdef,
                   Let("obj",rectyp, recrd,
                      Ident("obj"))) )
val t = typeOfProgram(TopLevel(p3))
assert(t == rectyp, s"Test failed: answer should be $rectyp")
passed(6)

//END TEST
```

Output:
```
Program type computed successfully as RecordType(List((set_counter,RefType(FunType(NumType,UnitType)))))

*** Tests Passed (6 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m(
  id = [32m"z"[39m,
  idType = NumType,
  e = [33mAssignRef[39m(e1 = [33mIdent[39m(s = [32m"x"[39m), e2 = [33mIdent[39m(s = [32m"z"[39m))
)
[36mgdef[39m: [32mFunDef[39m = [33mFunDef[39m(
  id = [32m"z"[39m,
  idType = UnitType,
  e = [33mDeRef[39m(e = [33mIdent[39m(s = [32m"x"[39m))
)
[36mrecrd[39m: [32mRecord[39m = [33mRecord[39m(ies = [33mList[39m(([32m"set_counter"[39m, [33mNewRef[39m(e = [33mIdent[39m(s = [32m"f"[39m)))))
[36mrectyp[39m: [32mRecordType[39m = [33mRecordType[39m(
  ies = [33mList[39m(([32m"set_counter"[39m, [33mRefType[39m(t = [33mFunType[39m(t1 = NumType, t2 = UnitType))))
)
[36mp3[39m: [32mLet[39m = [33mLet[39m(
  x = [32m"x"[39m,
  xType = [33mRefType[39m(t = NumType),
  e1 = [33mNewRef[39m(e = [33mConst[39m
```

<!-- grade_id: cell-fff1320a6c8f0f04, points: 6, grade: true, solution: false -->
```scala
//BEGIN TEST
/*
let x: ref(Num) = NewRef(0) in 
 let f : num => unit = function (z:num) { x := z} in 
  let g : unit => num = function (z:unit) {DeRef(x)} in
   let obj: {set_counter:num => unit, get_counter:unit=>num} = {set_counter = f, get_counter=g} in
     obj
*/
val fdef = FunDef("z", NumType, AssignRef(Ident("x"), Ident("z")))
val gdef = FunDef("z", UnitType, DeRef(Ident("x")))
val recrd = Record(List(("set_counter",Ident("f")), ("get_counter",Ident("g"))))
val rectyp = RecordType(List(("set_counter",FunType(NumType,UnitType)), 
                             ("get_counter",FunType(UnitType,NumType))))
val p3 = Let("x", RefType(NumType), NewRef(Const(0.0)), 
             Let("f", FunType(NumType, UnitType), fdef,
                Let("g", FunType(UnitType, NumType), gdef,
                   Let("obj",rectyp, recrd,
                      Ident("obj")))) )
val t = typeOfProgram(TopLevel(p3))
assert(t == rectyp, s"Test failed: answer should be $rectyp")
passed(6)

//END TEST
```

Output:
```
Program type computed successfully as RecordType(List((set_counter,FunType(NumType,UnitType)), (get_counter,FunType(UnitType,NumType))))

*** Tests Passed (6 points) ***
[36mfdef[39m: [32mFunDef[39m = [33mFunDef[39m(
  id = [32m"z"[39m,
  idType = NumType,
  e = [33mAssignRef[39m(e1 = [33mIdent[39m(s = [32m"x"[39m), e2 = [33mIdent[39m(s = [32m"z"[39m))
)
[36mgdef[39m: [32mFunDef[39m = [33mFunDef[39m(
  id = [32m"z"[39m,
  idType = UnitType,
  e = [33mDeRef[39m(e = [33mIdent[39m(s = [32m"x"[39m))
)
[36mrecrd[39m: [32mRecord[39m = [33mRecord[39m(
  ies = [33mList[39m(([32m"set_counter"[39m, [33mIdent[39m(s = [32m"f"[39m)), ([32m"get_counter"[39m, [33mIdent[39m(s = [32m"g"[39m)))
)
[36mrectyp[39m: [32mRecordType[39m = [33mRecordType[39m(
  ies = [33mList[39m(
    ([32m"set_counter"[39m, [33mFunType[39m(t1 = NumType, t2 = UnitType)),
    ([32m"get_counter"[39m, [33mFunType[39m(t1 = UnitType, t2 = NumType))
  )
)
[36mp3[3
```

## That's all folks!